---
title       : Insert the chapter title here
description : Insert the chapter description here
attachments :
  slides_link : https://s3.amazonaws.com/assets.datacamp.com/course/teach/slides_example.pdf

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:a540eafd9e
## A really bad movie

Give the correct answer. There's no hint!

*** =instructions
- Adventure
- Action
- Animation
- Comedy

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}
test_mc(2)
```

--- type:NormalExercise key:b9377ebc82
## Normal Exercise

asdf

*** =instructions
- asdf

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}
x <- 234
```

*** =solution
```{r}
x <- 234
```

*** =sct
```{r}
success_msg('test')
```
