---
title       : Testing PlainMultipleChoiceExercise
description : Testing PlainMultipleChoiceExercise

--- type:PlainMultipleChoiceExercise xp:50
## Plain Multiple Choice Exercise

Who was one of the creaters of R?

*** =instructions
- Ronald McDonald
- Ronald Reagan
- Donald Knuth
- Ross Ihaka


*** =sct
```{r}
test_mc(2)
```
